<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

# 1 Convertir un entier relatif
## 1.1 La mèthode du complément à deux.
### 1.1.1 Ecrire un nombre binaire sur X bits
Pour coder un entier relatifs en binaire avec la méthode du complément à deux, il est obligatoire de définir le nombre de bits servant à coder ce nombre.

On utilise habituellement 8, 16, 32 ou 64 bits.

**Ecrire une fonction `nbr_bits(chaine,n)` qui prend en arguments, une chaine de caractères représentant un nombre binaire et un entier qui représente le nombre bits sur lequel il faut coder ce nombre.**
> On ajoutera des zéros à gauche pour compléter `01011` doir renvoyer `00001011` pour une écriture sur **8 bits**.
> Si la chaine contient plus de bits que n, la fonction doit tronquer le nombre en gardant les caractères les plus à droite  par exmple `1011011101` sur **8 bits** doit renvoyer `11011101`.


<!-- grade_id: cell-9d1e62804aefa090, points: 0, grade: false, solution: true -->
```python
def nbr_bits(chaine,n) :
    # un nombre binaire sur X bits : by jaures Ratsim
    le = len(chaine)
    if chaine =="":
        return "0"*n
    return "0"*(n-le) + chaine if le<n else chaine[(le-n):le] if le>n else chaine
```

#### Jeu d'essai de la fonction `nbr_bits(chaine,n)`

<!-- grade_id: cell-506ed3301e5f7232, points: 2, grade: true, solution: false -->
```python
assert nbr_bits('10101001110',16) == '0000010101001110'
assert nbr_bits('1001011111111',8) == '11111111'
assert nbr_bits('10101010',8) == "10101010"
assert nbr_bits('', 32) == '00000000000000000000000000000000'
```

### 1.1.2 Inverser les bits.
La seconde étape de la méthode du complément à deux est d'inverser tous les bits.

**Ecrire une fonction `inverse_bit(chaine)` qui inverse chaque bit de la chaine de caractères. Un 1 devient un 0 et un 0 devient un 1**
>par exemple `10101` devient `01010`

<!-- grade_id: cell-05d25ef043683314, points: 0, grade: false, solution: true -->
```python
def inverse_bit(chaine) : 
    # initializing mapping dictionary
    map_dict = {'0':'1', '1':'0'}
    res = ''.join(i if i not in map_dict else map_dict[i] for i in chaine)
    return str(res)
```

#### Jeu d'essais de la fonction `inverse_bit(chaine)`

<!-- grade_id: cell-1da2576f857bc9d0, points: 1, grade: true, solution: false -->
```python
assert inverse_bit('111') == '000'
assert inverse_bit('0') == '1'
assert inverse_bit('10111101110110') == '01000010001001'
```

### 1.1.3 Additioner deux nombres binaires.

**On considére deux nombres binaires, vous devez écrire une fonction `add_binaire(chaine1,chaine2)`, cette fonction prend en paramètre deux chaines de caractères représentant un nombre écrit en base 2, et retourne une chaine de caractères représentant la somme de ces deux nombres sous forme binaire.**

> On commencera par mettre les deux chaines à la même taille, on pourra s'aider de la fonction `nbr_bits(n)`

> Pour commencer il sera utile d'inverser les chaines ce qui peut se faire avec : `chaine = chaine[::-1]`

> Enfin il suffira de parcourir tout les rangs des chaines en tenant éventuellement compte d'une retenue du rang précédent.


<!-- grade_id: cell-ec56d3763f9a8cb9, points: 0, grade: false, solution: true -->
```python
def add_binaire(a,b) :
    if int(a)==0 and int(b)==0  :
        return '0'*len(a) if (len(a)>len(b)) else '0'*len(b)
    return bin(int(a,2) + int(b,2))[2:] if (isinstance(a,str) and isinstance(b,str)) else None
```

#### Jeu d'essais de la fonction `add_binaire(chaine1,chaine2)`

<!-- grade_id: cell-ee630b125146e40e, points: 3, grade: true, solution: false -->
```python
assert add_binaire('1','1111') == '10000'
assert add_binaire('100','1') == '101'
assert add_binaire('0','00') == '00'
assert add_binaire('11','00') == '11'
```

## 1.2 Un programme de conversion pour entier relatif.
### 1.2.1 Une conversion de décimal vers binaire.
**Nous allons commencer par coder une fonction `dectobin(n)` qui transforme un entier positif ou nul exprimé en décimal, en une chaine de caractères qui code le nombre binaire correspondant.** 

<!-- grade_id: cell-4601cbfd9b19e220, points: 0, grade: false, solution: true -->
```python
def dectobin(n) :
    # Une conversion de décimal vers binaire
    return bin(n)[2:] if (isinstance(n,int)) else None
```

#### jeu d'essai de la fonction `dectobin(n)`

<!-- grade_id: cell-fca93c2e418d7366, points: 2, grade: true, solution: false -->
```python
assert dectobin(0) == '0'
assert dectobin(2) == '10'
assert dectobin(255) == '11111111'
assert dectobin(149) == '10010101'
assert dectobin(113) == '1110001'
```

### 1.2.2 conversion d'un entier relatif par la méthode du complément à 2
**En vous aidant des fonctions définies précédemment, ecrire une fonction de conversion `convert_entier(nbre,n)` qui transforme un entier relatif `nbre` (positif ou négatif) en une chaine de caractères codée sur `n` bits représentant ce nombre en binaire par la méthode du complément à 2.**

<!-- grade_id: cell-e62a18315ebc5c06, points: 0, grade: false, solution: true -->
```python
def convert_entier(nbre,n) :
    # YOUR CODE HERE
    raise NotImplementedError()
```

#### Jeu d'essais de la fonction `convert_entier(nbre,n)`

<!-- grade_id: cell-a267e076359495e3, points: 2, grade: true, solution: false -->
```python
assert convert_entier(-12,8) ==  '11110100' 
assert convert_entier(-12,16) ==  '1111111111110100'
assert convert_entier(256,16) ==  '0000000100000000'
assert convert_entier(-256,16) == '1111111100000000'
```

Output:
```
<ipython-input-2-0b8eea52b9a5> in <module>
----> 1 assert convert_entier(-12,8) ==  '11110100'
      2 assert convert_entier(-12,16) ==  '1111111111110100'
      3 assert convert_entier(256,16) ==  '0000000100000000'
      4 assert convert_entier(-256,16) == '1111111100000000'

<ipython-input-1-d508ce756bed> in convert_entier(nbre, n)
      1 def convert_entier(nbre,n) :
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

**Félicitations** vous venez d'écrire une fonction qui permet de coder n'importe quel entier relatif... ah c'est vrai il existe une façon simple et presque toute prête mais pas facile à expliquer pour l'instant sans un cours complet sur le formattage des chaines de caractères etl'utilisation de la bibliothèque struct, cepandant vous pouvez essayez le code ci-dessous :

```python
from struct import *
nombre = (''.join('{:08b}'.format(b) for b in pack('>b', -12)))
print("-12 sur 8 bits : ",nombre)
nombre = (''.join('{:08b}'.format(b) for b in pack('>h', -256)))
print("-256 sur 16 bits : ", nombre)
nombre = (''.join('{:08b}'.format(b) for b in pack('>i', 15746)))
print(" 15746 sur 32 bits : ", nombre)
```