<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

# 1. Un convertisseur décimal vers hexadécimal
## 1.1 La fonction de conversion.
On considére la fonction de conversion `dectohex(n)` codée ci-dessous, celle-ci prend en argument un nombre entier `n` et renvoie une chaine de caractère codant le nombre `n`.
**Complétez le texte ci-dessous en indiquant le type de chaque variable utilisées dans la fonction**
- n :
- hexa : 
- liste :

<!-- grade_id: cell-afb97d8c38e0174c, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

**Ce Script fonctionne parfaitement quelque soit `n` sauf quand `n=0`, modifiez le pour qu'il renvoie la chaine de caractères `"0"` quand le paramètre passée est `0`**


<!-- grade_id: cell-e10d8dd99b55dd93, points: 0, grade: false, solution: true -->
```python
def dectohex(n) :
    hexa = ""
    liste="0123456789ABCDEF"
    if (n==0):
        return "0"
    else:
        while n != 0 :
            hexa = liste[n%16] + hexa
            n = n // 16
        return(hexa)
```

Output:
```
'0'
```

**Jeu de tests :**

<!-- grade_id: cell-4f9b8cfcc84237b9, points: 2, grade: true, solution: false -->
```python
assert dectohex(255) == 'FF'
assert dectohex(16) == '10'
assert dectohex(0) == '0'
assert dectohex(256) == '100'
```

# 2 Formatage d'un nombre héxadécimal en bloc d'octet.
Un octet rerprésente 8 bits, par exemple `10100110`, ce nombre d'écrit `A6` en hexadécimal.

**Ecrire une fonction `octet(n)` qui scinde la chaine de caractères passée en argument qui représente un nombre codé en hexadécimal en une chaines de caractéres représentant plusieurs octets séparés par des espaces**

> par exemple `octet('A6B518') ` doit renvoyer : `'A6 B5 18'`.

<!-- grade_id: cell-9695d5647e1135a9, points: 0, grade: false, solution: true -->
```python
def octet(n):
    f,L,sep = (len(n)),[], ' '
    if (f%2 !=0):
        n = "0"+n
    ["00" if n =="0" else [L.append(n[i-1:i+1]) for i in range(1,f+1,2)]]
    return sep.join(L)
```

<!-- grade_id: cell-00f8519b3e90965e, points: 2, grade: true, solution: false -->
```python
assert octet("AD2B6547") == 'AD 2B 65 47'
assert octet("2B5C7") == '02 B5 C7'
assert octet("0") == "00"
```

```python
def octet(n):
    # Defining splitting point
    sep,j = " ",2
    # Using list comprehension
    L = [(n[i:i+j]) for i in range(0, len(n), j)]
    return sep.join(L)
```