<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

# 1. Et les shadoks comptaient
## 1.1 La numération shadok
Rendez vous [sur la vidéo numération shadok](https://www.youtube.com/watch?v=lP9PaDs2xgQ) pour mieux comprendre comment comptent les shadoks.
>**L'objectif de cette activité est d'aider les shadoks à cosntruire leur première calculatrice... une calculatrice qui fait des additions des soustractions et même des multiplications.**

## 1.2 Un convertisseur shadok vers décimal
### 1.2.1. La fonction de conversion
Considérons la définition de la fonction ci-dessous, cette fonction prend une chaine de caractères en argument qui représente le nombre shadok, et renvoie l'entier correspondant.

```python
def shadoktodec(chaine):
    shadok =["GA" , "BU" , "ZO" , "MEU"] #shadok est une liste
    chaine = chaine.split(" ") # les différents caractères sont séparés par des espaces.     
    resultat = 0
    rang = 1
    for i in reversed(chaine) : # reversed permet de prendre la liste à partir du dernier rang.
        for j in range(4) :
            if i == shadok[j] :
                resultat = resultat + j*4**rang
                rang += 1
    return (resultat)
```

```python
shadoktodec("BU MEU")
```

Output:
```
28
```

Ce script ne passe le jeu de tests, il contient visiblement une erreur.. Les valeurs renvoyées ne correspondent pas ..
>**Modifier ce script ci dessous, pour qu'il passe le jeu de test**

<!-- grade_id: cell-899237a71093edab, points: 0, grade: false, solution: true -->
```python
def shadoktodec(chaine):
    shadok =["GA" , "BU" , "ZO" , "MEU"] #shadok est une liste
    chaine = chaine.split(" ") # les différents caractères sont séparés par des espaces.     
    resultat = 0
    rang = 0
    for i in reversed(chaine) : # reversed permet de prendre la liste à partir du dernier rang.
        for j in range(4) :
            if i == shadok[j] :
                resultat = resultat + j*4**rang
                rang += 1
    return (resultat)
```

### 1.2.2. Jeu d'essais 

<!-- grade_id: cell-3225f27ae655d4c3, points: 1, grade: true, solution: false -->
```python
assert shadoktodec("BU BU") == 5
assert shadoktodec("GA") == 0
assert shadoktodec("MEU ZO BU") == 57
assert shadoktodec("BU BU") == 5
assert shadoktodec("ZO BU") == 9
```

## 1.3 Un convertisseur décimal vers Shadok
### 1.3.1  La fonction de conversion.
Cette fonction prend un entier positif comme argument, et retourne une chaine de caractères donnant le nombre en shadok.
Ce script a un défaut,il renvoie une chaine de caractères se terminant par un espace.("BU GA ZO " par exemple)


```python
def dectoshadok(n) :
    shadok =["GA" , "BU" , "ZO" , "MEU"]
    resultat = ""
    while n > 4 :
        resultat =  shadok[n % 4] + " " + resultat 
        n = n // 4
    resultat =  shadok[n] + " " + resultat
    return (resultat)
```

Ce script ne passe pas les tests et renvoie même parfois une erreur...

```python
dectoshadok(1)
```

Output:
```
'BU '
```

```python
dectoshadok(2457)
```

Output:
```
'ZO BU ZO BU ZO BU '
```

> **Modifier ce script de telle façon qu'il passe tous les test du jeu de tests ci-dessous.**

<!-- grade_id: cell-d563d55f3ebb58c4, points: 0, grade: false, solution: true -->
```python
def dectoshadok(n) :
    shadok =["GA" , "BU" , "ZO" , "MEU"]
    resultat = ""
    while n >= 4 :
        resultat =  shadok[n % 4] + " " + resultat 
        n = n // 4
    resultat =  shadok[n] + " " + resultat
    return (resultat)
```

### 1.3.2 le jeu de tests.
Votre fonction doit passer l'ensemble du jeu de tests ci-dessous :

<!-- grade_id: cell-4bc2723da47902e9, points: 1, grade: true, solution: false -->
```python
assert dectoshadok(1) == "BU "
assert dectoshadok(16) == "BU GA GA "
assert dectoshadok(2457) == "ZO BU ZO BU ZO BU "
```

# 2. Une calculatrice shadok.
## 2.1 Consignes
Vous devez écrire une fonction `addshadock(chaine1,chaine2)` qui prend en argument deux chaines de caractères représentant un nombre en shadok et qui renvoie une chaine de caractères qui représente la somme des deux nombre passés en argument.
> Rq : il serait pertinent d'utiliser les fonctions précédemment écrites.

<!-- grade_id: cell-cba075c2508fb5f6, points: 0, grade: false, solution: true -->
```python
def addshadok(chaine1,chaine2) :
    dec1,dec2 = shadoktodec(chaine1),shadoktodec(chaine2)
    somDec =  dec1 + dec2
    return dectoshadok(somDec)
def shadoktodec(chaine):
    shadok =["GA" , "BU" , "ZO" , "MEU"] #shadok est une liste
    chaine = chaine.split(" ") # les différents caractères sont séparés par des espaces.     
    resultat = 0
    rang = 0
    for i in reversed(chaine) : # reversed permet de prendre la liste à partir du dernier rang.
        for j in range(4) :
            if i == shadok[j] :
                resultat = resultat + j*4**rang
                rang += 1
    return (resultat)
def dectoshadok(n) :
    shadok =["GA" , "BU" , "ZO" , "MEU"]
    resultat = ""
    while n >= 4 :
        resultat =  shadok[n % 4] + " " + resultat 
        n = n // 4
    resultat =  shadok[n] + " " + resultat
    return (resultat)
```

## 2.2 Jeu de tests
Votre fonction doit passer le jeu de tests suivant :

<!-- grade_id: cell-37e3b20b329d70f6, points: 2, grade: true, solution: false -->
```python
assert addshadok("GA","GA") == 'GA '
assert addshadok("BU","GA") == 'BU '
assert addshadok("BU","MEU") == 'BU GA '
assert addshadok("BU BU","MEU GA GA") == 'MEU BU BU '
```

```python

```