<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

## 1 Un convertisseur héxadécimal vers décimal
On considère la fonction `hextodec(chaine)` partiellement programmée ci dessous, cette fonction prend en argument un chaine de caractéres représentant un nombre codé en hexadécimal, et renvoie un nombre en décimal sous la forme d'un entier.

**Le code de cette fonction n'est pas complet, compléter le pour que la fonction puisse passer le jeu d'essais.**

<!-- grade_id: cell-2858bb3ab4eac687, points: 0, grade: false, solution: true -->
```python
def hextodec(chaine) :
    return int(chaine, 16) if (isinstance(chaine,str)) else None
```

**Jeu d'essais**

<!-- grade_id: cell-1fa723eb39066510, points: 2, grade: true, solution: false -->
```python
assert hextodec('FF') == 255
assert hextodec('0') == 0
assert hextodec('10') == 16
assert hextodec('A5B4287F') == 2780047487
```

# 2. Un convertisseur hexadécimal vers binaire.
On souhaite obtenir une fonction `hextobin(chaine)` qui prend en argument une chaine de caractère et qui renvoie une chaine de caractères représentant ce nombre exprimé en binaire.
## 2.1 Convertisseur décimal vers binaire
**Pour cela , commençons par écrire une fonction `dectobin(n)` qui prend un nombre décimal en argument et qui va renvoyer une chaine de caractères exprimant ce même nombre en binaire :**


<!-- grade_id: cell-a845d6a98937bb76, points: 0, grade: false, solution: true -->
```python
def dectobin(n) :
    # YOUR CODE HERE
    return bin(n)[2:] if (isinstance(n,int)) else None
```

**Jeu d'essais**

<!-- grade_id: cell-6fe2a254f102f177, points: 2, grade: true, solution: false -->
```python
assert dectobin(0) == '0'
assert dectobin(4) == '100'
assert dectobin(255) == '11111111'
```

## 2.2 Convertisseur hexadécimal binaire.
En utilisant les deux fonctions précédemment écrites, on veut maintenant créer un convertisseur hexadécimal/Binaire.

**Ecrire une fonction `hextobin(chaine)` qui prend en argument une chaine de caractères représentant un nombre codé en hexadécimal et qui renvoie une chaine de caractères représentant ce même nombre codé en binaire.**

<!-- grade_id: cell-0f39b0f7a6bbbc4a, points: 0, grade: false, solution: true -->
```python
def hextobin(chaine) :
    # YOUR CODE HERE
    return bin(int(chaine, 16))[2:] if (isinstance(chaine,str)) else None
    
```

### Jeu d'essai

<!-- grade_id: cell-ee55bb5ec47311b8, points: 2, grade: true, solution: false -->
```python
assert hextobin('FF') == '11111111'
assert hextobin('0') == '0'
assert hextobin('A15FD3') == '101000010101111111010011'
```

## 2.3 Ecriture du nombre en bloc d'un Octet.
Pour faciliter la mesure des nombres binaires, il peut être pratique de les écrire par bloc de 8 bits, c'est à dire d'un octet, par exemple plutôt que d'écrire `10110101011010101010101` on préférera écrire `1011010 10110101 01010101`.

Si le nombre de bits n'est pas un multiple de 8, on rajoutera des 0 à gauche pour compléter, par exemple `101101010010`, on écrira `00001011 01010010`

**Ecrire une fonction `octet(chaine)`, qui prend en argument une chaine de caractères représentant un nombre en binaire, et la formate en bloc d'octet, conformément à la description ci-dessus**


<!-- grade_id: cell-46a2e4ce087d8d8c, points: 0, grade: false, solution: true -->
```python
def octet(n) :
    # Defining splitting point
    sep,j,x = " ",8,len(n)
    # Using list comprehension
    if x%8 != 0:
        n = "0"*(8-x%8) + n 
    L = [(n[i:i+j]) for i in range(0, len(n), j)]
    return sep.join(L)
```

**Jeu d'essai** 

<!-- grade_id: cell-7e19625003bb7130, points: 0, grade: true, solution: false -->
```python
assert octet("11110001") == '11110001'
assert octet("1001110001") == '00000010 01110001'
assert octet("0") == '00000000'
assert octet("10101101101101101101110111011101110") == '00000101 01101101 10110110 11101110 11101110'
```

```python

```