<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

# 2. Du décimal et binaire
## 2.1 Généralités
Les nombres que nous utilisons sont des nombres dits pondérés dans une base. Dans chaque base, nous utilisons un symbole distinct pour chaque chiffre. Chaque chiffre appartient à un rang. Ce rang est pondéré. On peut donc généraliser la forme d’écriture d’un nombre ainsi :  $ c_{i-1} \cdot b_{i-1} + ... + c_{2} \cdot b_{2} + c_{1} \cdot b_{1} + c_{0} \cdot b_{0} $

Où $c_{i}$  est le chiffre au rang i de la base b. Cela s’écrit aussi sous cette forme : $$ \sum_{i=1}^{n} c_{i} \cdot b_{i}$$

## 2.2 Les numérations antiques (sans chronologie...)

### 2.2.1 Numération Egyptienne (A titre culturel)

C’est une numération **additionnelle** : le résultat est obtenu en additionnant les chiffres sans tenir compte de leurs positions.

![hiéroglyphes](Images\egyp1.png)

[source image](http://jlsigrist.com/egyptien.html)

### 2.2.2 Numération romaine (A titre culturel)

![Chiffres romains](Images\romain1.jpg)

[source image](http://villemin.gerard.free.fr/Wwwgvmm/Numerati/Romain.htm)

![Exemples romain](Images\romain2.jpg)

C’est aussi une numération **additionnelle** : le résultat est obtenu en additionnant les chiffres. Mais cette fois les positions ont de l'importance. (XIX et XXI)

Remarquez-vous l'absence d'un chiffre important pour nous ?

### 2.2.3 Numération mésopotamienne (A titre culturel)

C'est une numération **positionnelle** en **base 60**.  Il en reste quelques vestiges dans le système horaire ou dans la mesure des angles en degrés, minutes, secondes.

![Mésopotamie](Images\base_60.gif)


[source image](http://www.lemiroirdutemps.com/article-212-les-nombres-sont-ils-magiques-il-y-a-environ-6-000-ans-38724934.html)


## 2.3 La numération binaire (base 2)

### 2.3.1 Théorie

On va donc avoir besoin de 2 chiffres. Lesquels ?	GA	puis	BU ?

Non on va faire simple 0 puis 1. On aura donc que 2 chiffres.

![Binaire](Images\bin_101.png)

On a ici un nombre écrit sur 3 bits.

​		Le bit 1 (à droite) se trouve en **position 0** (2 puissance 0 = 1), il a le poids le plus faible (1).

​	Le bit 0 (au milieu) se trouve en **position 1** (2 puissance 1 = 2). Il a le poids 2

Le bit 1 (à gauche) se trouve en **position 2** (2 puissance 2 = 4), il a le poids le plus fort (4).



On les prends par paquet de combien ? **8** c'est pas mal ? On dira qu'on prend **un octet**

On a donc les poids suivants pour chaque bit :

| Position        | 7    | 6    | 5    | 4    | 3    | 2    | 1    | 0    |
| --------------- | :----: | :----: | :----: | :----: | :----: | :----: | :----: | :----: |
| - | 2^7 | 2^6 | 2^5 | 2^4 | 2^3 | 2^2 | 2^1 | 2^0 |
| Poids en base 2 | 128  | 64 | 32 | 16 | 8 | 4 | 2 | 1 |

Comment connaître le poids de la colonne 6 ? C'est simple c'est le double de la colonne 5 !

*remarque :*

> 8, 16, 32, 64, mais surtout pas 48... Ces chiffres vous disent quelque chose ? Ce sont les capacités possibles pour la mémoire des téléphone, carte mémoire ou  clé USB : c'est forcément en 2^??

### 2.3.2 EXERCICE : un convertisseur binaire vers décimal
A partir de ce que vous avez fait avec les Shadoks, écrire une fonction `bintodec(chaine)` qui converti un binaire (donné en paramètre **sous forme de chaine de caractère**, "01001010") en décimal, sous forme de nombre entier.

<!-- grade_id: cell-dc0477901f0b86a1, points: 0, grade: false, solution: true -->
```python
def bintodec(chaine) :
    rang = len(chaine) - 1
    resultat = 0
    for i in range (len(chaine)):
        if chaine[i] == "1":
            resultat += 2 ** rang
        rang -= 1
    return resultat
#En local ça ne marche pas sans le mot clé : "print"!!!
bintodec ("1111")
```

Output:
```
15
```

Voyons maintenant si votre fonction passent les tests : 

<!-- grade_id: cell-a06a24882d66b53d, points: 1, grade: true, solution: false -->
```python
assert bintodec('10') == 2
assert bintodec('1111') == 15
assert bintodec('0000') == 0
assert bintodec('10000000') == 128
assert bintodec('1') == 1
```

### 2.3.3 EXERCICE : un convertisseur décimal vers binaire.
A partir de ce que vous avez fait avec les Shadoks, écrire une fonction `dectobin(n)` qui converti un nombre entier `n` passé en paramètre en un binaire (donné **sous forme de chaine de caractère**, "01001010" par exemple).

<!-- grade_id: cell-383e2f0fe6d29697, points: 0, grade: false, solution: true -->
```python
def dectobin(n):
    resultat = ""
    while n >= 2 :
        resultat =  str(n % 2) + resultat 
        n = n // 2
    resultat =  str(n) + resultat
    return (resultat)
```

Voyons maintenant si votre question passe le jeu de tests.

<!-- grade_id: cell-e53241952466fd89, points: 1, grade: true, solution: false -->
```python
assert dectobin(7) == '111'
assert dectobin(1) == '1'
assert dectobin(1060) == '10000100100'
assert dectobin(0) == '0'
```

### 2.3.4 EXERCICE : Combien de bit pour écrire un nombre en binaire ?
Si on considère le nombre écrit en binaire '110', on constate qu'il faut 3 bits pour l'écrire, alors qu'il en faut 8 pour écrire '11101010'.

Ecrire une fonction `nbrbit(n)`, qui prend en argument un entier positif ou nul et qui renvoie le nombre de bits necéssaires pour écrire ce nombre en binaire.

<!-- grade_id: cell-7dc062cd1b496c33, points: 0, grade: false, solution: true -->
```python
def nbrbit(n) :
    return len(list(dectobin(n)))
def dectobin(n):
    resultat = ""
    while n >= 2 :
        resultat =  str(n % 2) + resultat 
        n = n // 2
    resultat =  str(n) + resultat
    return (resultat)
```

Voyons voir le jeu de tests :

<!-- grade_id: cell-e6c8f6f4f7204bbb, points: 2, grade: true, solution: false -->
```python
assert nbrbit(1) == 1
assert nbrbit(15) == 4
assert nbrbit(16) == 5
assert nbrbit(0) == 1
```

```python

```