<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

<div align=justify>
    <h1 style='text-align:center' ><span style = 'color:#0000FF'>NoteBook : Epreuve pratique n°1 </span></h1>
</div>

# Exercice 1 (8 points)

<div class = "alert alert-info">
    Écrire en langage Python une fonction <code>recherche</code> prenant comme paramètres une variable <code>a</code> de type numérique (float ou int) et un tableau <code>t</code> (type list) et qui renvoie le nombre d'occurrences de <code>a</code> dans <code>t</code>.
</div>

<!-- grade_id: cell-11ae7f6d12b1f306, points: 0, grade: false, solution: true -->
```python
def recherche (a,t) :
    nb_caracteres = 0  # Variable stockant le nombre de caractères trouvé, initialisée à 0
    for c in t:  # On parcourt chaque caractère du mot passé en argument
        if c == a:  # Si le caractère courant est égal au caractère recherché
            nb_caracteres += 1  # On incrémente nb_caracteres
    return nb_caracteres  # Enfin, on retourne le nombre de caractères trouvé

```

<!-- grade_id: cell-1ea587ee5da7d980, points: 8, grade: true, solution: false -->
```python
assert recherche(5,[]) == 0
assert recherche(5,[-2, 3, 4, 8]) ==0
assert recherche(5,[-2, 3, 1, 5, 3, 7, 4]) == 1
assert recherche(5,[-2, 5, 3, 5, 4, 5]) == 3
```

# Exercice 2 (12 points)

<div class = "alert alert-info">
    La fonction rendu_monnaie_centimes prend en paramètres deux nombres entiers positifs <code>s_due</code> et <code>s_versee</code> et elle permet de procéder au rendu de monnaie de la différence <code>s_versee – s_due</code> pour des achats effectués avec le système de pièces de la zone Euro. 
    
    On utilise pour cela un algorithme qui commence par rendre le maximum de pièces de plus grandes valeurs et ainsi de suite. La fonction renvoie la liste des pièces qui composent le rendu.
    
Toutes les sommes sont exprimées en centimes d’euros. Les valeurs possibles pour les pièces sont donc <code>[1, 2, 5, 10, 20, 50, 100, 200]</code>.
    
Ainsi, l’instruction <code>rendu_monnaie_centimes(452, 500)</code> renverra la liste <code>[20, 20, 5, 2, 1]</code>.
    En effet, la somme à rendre est de <b>48</b> centimes soit <b>20 + 20 + 5 + 2 + 1</b>.
</div>

Le code de la fonction est donné ci-dessous :

```python
def rendu_monnaie_centimes(s_due, s_versee):
    pieces = [1, 2, 5, 10, 20, 50, 100, 200]
    rendu = ...
    a_rendre = ...
    i = len(pieces) - 1
    while a_rendre > ... :
        if pieces[i] <= a_rendre :
            rendu.append(...)
            a_rendre = ...
        else :
            i = ...
    return rendu
```

<div class = "alert alert-info">
    Recopier et modifier le code ci-dessus pour qu'il passe les assertions
<div>

<!-- grade_id: cell-1abf4ec7da0e54f7, points: 0, grade: false, solution: true -->
```python
def rendu_monnaie_centimes(s_due, s_versee):
    pieces = [1, 2, 5, 10, 20, 50, 100, 200]
    rendu = []
    a_rendre = s_versee - s_due
    i = len(pieces) - 1
    while a_rendre > 0:
        if pieces[i] <= a_rendre :
            rendu.append(pieces[i])
            a_rendre = a_rendre - pieces[i]
        else :
            i = i-1
    return rendu
```

<!-- grade_id: cell-7873874ee9593ed7, points: 8, grade: true, solution: false -->
```python
assert rendu_monnaie_centimes(700,700) == []
assert rendu_monnaie_centimes(112,500) == [200, 100, 50, 20, 10, 5, 2, 1]
assert rendu_monnaie_centimes(112,120) == [5, 2, 1]
assert rendu_monnaie_centimes(51,120) == [50, 10, 5, 2, 2]
```

<div class = "alert alert-info">
 <b>Q1</b> : A quelle famille d'algorithme appartient l'algorithme ci-dessus appartient-il ?
</div>

<!-- grade_id: cell-d11b7821e5f43d82, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE Glouton???

<div class = "alert alert-info">
 <b>Q2</b> : La solution renvoyée par ce type d'algorithme n'est pas toujours optimale, pourtant le système de rendu de monnaie utilisé ici garanti dans ce cas une solution optimal. Comment s'appelle un tel système de rendu de monnaie ?
</div>

<!-- grade_id: cell-4c5c844c77be4f26, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE : canonique

<div class="alert alert-warning" align="center">
    <b>Equipe pédagoqique LFT-NSI</b>
<br/>ressource éducative libre distribuée sous <a href="http://creativecommons.org/licenses/by-nc-sa/4.0/"!>Licence Creative Commons Attribution</a>
    <img src="https://i.creativecommons.org/l/by-nc-sa/4.0/88x31.png" align="center" width="70"  margin="50px"> 
</div>